# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

```python
import pandas as pd
df = pd.read_csv("assets/NISPUF17.csv")
df["EDUC1"].value_counts()
```

Output:
```
4    13656
3     6999
2     4906
1     2904
Name: EDUC1, dtype: int64
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    edu = df["EDUC1"].value_counts(normalize=True)

    return {
        "less than high school": edu.loc[1],
        "high school": edu.loc[2],
        "more than high school but not college": edu.loc[3],
        "college": edu.loc[4]
    } 
    raise NotImplementedError()
    
proportion_of_education()
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
# CBF_01 refers to the categories of breastfed babies
# 1=yes, 2=no, 3=don't know, 4=missing
df["CBF_01"].value_counts()
```

Output:
```
1     24261
2      4085
77      110
99        9
Name: CBF_01, dtype: int64
```

```python
# P_NUMFLU is the total column for all types of seasonal influenza shots the child had
df["P_NUMFLU"]
```

Output:
```
0        NaN
1        NaN
2        NaN
3        3.0
4        0.0
        ... 
28460    NaN
28461    NaN
28462    NaN
28463    NaN
28464    NaN
Name: P_NUMFLU, Length: 28465, dtype: float64
```

```python
# We need to map the conditions of a breastfed baby to the CBF_01 column and then calculate the mean value from the true numbers using the P_NUMFLU column

bf_avg = df[(df["CBF_01"] == 1)]["P_NUMFLU"].mean()
bf_avg
```

Output:
```
1.8799187420058687
```

```python
# We need to map the conditions of a non-breastfed baby to the CBF_01 column and then calculate the mean value from the true numbers using the P_NUMFLU column

nbf_avg = df[df["CBF_01"] == 2]["P_NUMFLU"].mean()
nbf_avg
```

Output:
```
1.5963945918878317
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
# Now combining what we have learned in the previous cells:

def average_influenza_doses():

    return (df[(df["CBF_01"] == 1)]["P_NUMFLU"].mean(),
           df[df["CBF_01"] == 2]["P_NUMFLU"].mean())

    raise NotImplementedError()
    
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
# First check to see if the SEX column provides us with the 2 genders:
# 1=male, 2=female

df["SEX"].value_counts()
```

Output:
```
1    14717
2    13748
Name: SEX, dtype: int64
```

```python
# Varicella doses are denoted in the P_NUMVRC column

df["P_NUMVRC"]
```

Output:
```
0        NaN
1        NaN
2        NaN
3        1.0
4        0.0
        ... 
28460    NaN
28461    NaN
28462    NaN
28463    NaN
28464    NaN
Name: P_NUMVRC, Length: 28465, dtype: float64
```

```python
# Find categories in column HAD_CPOX:
# 1=yes, 2=no, 77=don't know, 99=missing

df["HAD_CPOX"].value_counts()
```

Output:
```
2     27955
1       402
77      105
99        3
Name: HAD_CPOX, dtype: int64
```

```python
# Calculate how many males had the vaccine and were still infected

m_hadcp = (df[(df["SEX"] == 1) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 1).sum()
m_nhadcp = (df[(df["SEX"] == 1) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 2).sum()
m_ratio = m_hadcp/m_nhadcp

print("Male infected with cpox after vaccinated: {}".format(m_hadcp))
print("Male ratio of {}".format(m_ratio))

# Calculate how many females had the vaccine and were still infected
f_hadcp = (df[(df["SEX"] == 2) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 1).sum()
f_nhadcp = (df[(df["SEX"] == 2) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 2).sum()
f_ratio = f_hadcp/f_nhadcp

print("Female infected with cpox after vaccinated: {}".format(f_hadcp))
print("Female ratio of {}".format(f_ratio))
```

Output:
```
Male infected with cpox after vaccinated: 68
Male ratio of 0.009675583380762664
Female infected with cpox after vaccinated: 53
Female ratio of 0.0077918259335489565
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    return {
        "male": (df[(df["SEX"] == 1) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 1).sum() / (df[(df["SEX"] == 1) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 2).sum(),
        "female": (df[(df["SEX"] == 2) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 1).sum() / (df[(df["SEX"] == 2) & (df["P_NUMVRC"] > 0.0)]["HAD_CPOX"] == 2).sum()
    }
      
    raise NotImplementedError()
    
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
df["P_NUMVRC"].info()
```

Output:
```
<class 'pandas.core.series.Series'>
RangeIndex: 28465 entries, 0 to 28464
Series name: P_NUMVRC
Non-Null Count  Dtype  
--------------  -----  
15333 non-null  float64
dtypes: float64(1)
memory usage: 222.5 KB
```

```python
df["HAD_CPOX"].info()
```

Output:
```
<class 'pandas.core.series.Series'>
RangeIndex: 28465 entries, 0 to 28464
Series name: HAD_CPOX
Non-Null Count  Dtype
--------------  -----
28465 non-null  int64
dtypes: int64(1)
memory usage: 222.5 KB
```

```python
df["P_NUMVRC"]
```

Output:
```
0        NaN
1        NaN
2        NaN
3        1.0
4        0.0
        ... 
28460    NaN
28461    NaN
28462    NaN
28463    NaN
28464    NaN
Name: P_NUMVRC, Length: 28465, dtype: float64
```

```python
df.dropna(subset=["P_NUMVRC"], inplace=True)
df
```

Output:
```
Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  \
3               4    86131      8613     1  675.430817  333.447418     1040   
4               5   227141     22714     1  482.617748  278.768063     1008   
6               7   221381     22138     1  516.427723  248.916038     1008   
7               8   182861     18286     1  313.225832  163.213401     1001   
8               9   161121     16112     1  978.042448  718.616990     1016   
...           ...      ...       ...   ...         ...         ...      ...   
28446       28447   193801     19380     1  321.058203  110.009181     2047   
28447       28448   136841     13684     1  593.426753  337.255316     2031   
28449       28450   132081     13208     1  166.681363   96.330007     2059   
28452       28453    82151      8215     1  335.456813  159.895243     2056   
28454       28455   191451     19145     1   59.626969   29.187316     2074   

       YEAR  AGECPOXR  HAD_CPOX  ...  XVRCTY2  XVRCTY3  XV
```

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import pandas as pd
    import scipy.stats as stats
    
    df = pd.read_csv("assets/NISPUF17.csv")
    
    # keep only valid responses (1 = yes, 2 = no)
    df = df[(df["HAD_CPOX"].isin([1, 2])) & (df["P_NUMVRC"].notna())]
    
    corr, _ = stats.pearsonr(df["HAD_CPOX"], df["P_NUMVRC"])
    
    return corr

corr_chickenpox()
```

Output:
```
0.07044873460147986
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```